Quantifying microbial fitness in high-throughput experiments

Author response:

The following is the authors’ response to the previous reviews

We thank both editors and the three reviewers for their positive feedback on the revised manuscript. Following suggestions from Reviewer #1, we cite additional literature to clarify the use of ’fitness potential’ and we have revised the caption of Figure 3 to better explain how we generate the panel of mutants. We have also added a sentence to emphasize that the selection coefficient should match the time-scale of bottleneck effects in the evolution environment.

Changes from the earlier version to the revised version

Abstract

Few concepts are as central to evolution as is fitness, and yet the quantification of fitness is often ambiguous. In particular, high-throughput experiments to measure mutant fitness in microbes are increasingly common but vary widely in their definitions of fitness, which makes their results difficult to compare. What are the consequences of these different fitness statistics, and is there a best way to quantify fitness in a given context? Here we systematize the set of possible fitness statistics according to the following three choices: 1) the encoding of relative abundance (e.g., transforming by a log or logit function), 2) the time scale over which to measure the change in relative abundance, and 3) the choice of a reference subpopulation for calculating relative fitness in bulk competition experiments, such as those using DNA-barcoded mutants. We show that these different choices can lead to significantly different interpretations of mutant fitness, affecting the magnitude of fitness effects, the apparent presence of epistasis, and even the fitness ranking across mutants. This can confound predictions for evolutionary dynamics and gene functions. Altogether our results demonstrate the importance of consistent fitness definitions for reproducible results across experiments.

Introduction

The concept of fitness is supposed to describe the fate of mutations as they arise in a population [1] and defines other core evolutionary concepts such as tradeoffs [2]. At the same time, fitness is often a confusing term in evolutionary biology [3–5]. It is possible to classify fitness measures by their role in theory [6] and their scale of measurement [7–9], but these arguments offer little guidance on how to measure fitness in practice. Empirical fitness measurements have become widespread in the microbial sciences — to test evolutionary theory [10–15], detect microbial interactions [16, 17], annotate gene functions [18, 19], and understand the spread of antibiotic resistance genes [20, 21] — but the approaches to quantifying fitness in these studies vary widely.

The classic approach to measuring the fitness effects of mutations uses pairwise competition experiments, where the mutant competes with its ancestral wild-type in co-culture [22–24]. This approach is ideal because it mimics the dynamics of spontaneous mutations, and thus is typically used in experimental evolution [10, 12, 25–27]. For example, the Long-Term Evolution Experiment (LTEE) has evolved Escherichia coli over tens of thousands of generations and measured the fitness of each evolved lineage in competition with the ancestor [11, 26, 28]. However, pairwise competition experiments require a way to distinguish mutants from the wild-type in cocultures (e.g., fluorescent labeling or selection markers), which is not always feasible to engineer. A second approach is thus to measure properties of genotypes in monocultures, like the growth rate, and estimate relative fitness between genotypes from these measurements. Monoculture growth curves are straightforward to measure, even for large numbers of mutants [29–31], but empirical tests have shown that the monoculture growth rate is insufficient to predict the relative fitness between genotypes [32–34]. Summary statistics of whole growth curves, like the area under the curve (AUC), perform better [35–37]. However, this approach breaks down for strains that engage in cross-feeding or toxin production, and attempts to improve monoculture-based fitness estimates require additional experiments [34, 37] and do not go beyond qualitative fitness rankings [37].

A third and more recent approach is to measure relative fitness in bulk competition experiments using high-throughput sequencing, often of DNA barcodes [13, 38]. For example, transposon-insertion mutagenesis generates libraries of gene knockouts that can be tracked as they grow in a single batch culture [39]. Estimating the fitness of each mutant allows us to identify genes that are particularly important in the growth environment [18, 19, 40], in the given genetic background [14, 15, 41, 42], and in the context of an ecological community [16, 17, 42, 43]. Bulk fitness estimation also plays a major role in identifying beneficial lineages in experimental evolution [13, 44] (not unlike the tracking of viral pathogens in public health [45–47]) and especially testing their fitness in new environments [48–52]. Many of these studies make different choices about the design of experiments and the calculation of fitness, without explaining why a certain fitness measure was used or how to compare between data sets.

The inconsistent choices for quantifying microbial fitness make it is difficult to compare fitness across experiments, especially because different groups tend to use different definitions. For example, the LTEE reports relative fitness per-generation [11, 26], whereas other evolution experiments exclusively report relative fitness percycle of batch growth [12, 53]. New techniques, like barcode sequencing, often spark their own fitness metrics [18, 40, 54–56]. A different choice of fitness metric might lead to a different ranking of mutants in a genetic screen (and the importance we assign to those genes) or change statistical properties of the distribution of mutant fitness effects, which is essential to predict the probability of mutant fixation [57], the speed of adaptation [58], and the effect size in statistical tests of gene function [18, 19]. This raises three key questions: 1) How do these fitness statistics differ? 2) Can we say which choices are optimal, in the light of first principles or some practical considerations? 3) Does this tell us how we should design these experiments? Previous work has addressed aspects of these questions [6–9, 59] but the arguments are scattered across disciplines and based on generic models of population growth [6, 7, 9, 60], rather than incorporating the specific dynamics of microbes in laboratory experiments. Here we address these questions from a unified framework, using realistic microbial population dynamics with empirical traits to systematically test different fitness metrics.

Results

Predictive power of fitness depends on the encoding of relative abundance

Earlier version: We distinguish between three related but distinct notions of fitness [6], all based on predicting dynamics of a population [61–64]. The first type of fitness is absolute fitness, which is a property of a single genotype by itself and serves to predict the change in the genotype’s absolute abundance N (t) over a future time window Δt (Figure S1A). This is important for questions about extinction and evolutionary rescue [65]. The second type of fitness is relative fitness, which is a property of two genotypes as it describes how the relative abundance x(t) of one genotype changes compared to the other over a time Δt (Figure S1B). This is important to determine the fixation probability of new mutations [57, 66, 67]. In general these dynamics are stochastic [66, 68, 69], but throughout this paper we focus on their average behavior across replicate cultures (as sketched in Figure S1C).

Revised version: We distinguish between three related but distinct notions of fitness [6], all based on predicting dynamics of a population [61–64]. The first type of fitness is absolute fitness, which is a property of a single genotype by itself and serves to predict the change in the genotype’s absolute abundance N (t) over a future time window Δt (Fig. S1A). This is important for questions about extinction and evolutionary rescue [65]. The second type of fitness is relative fitness, which is a property of two geno-types as it describes how the relative abundance x(t) of one genotype changes compared to the other over a time Δt (Fig. S1B). This is important to determine the fixation probability of new mutations [57, 66, 67]. In general these dynamics are stochastic [66, 68, 69], but throughout this paper we focus on their average behavior across replicate cultures (as sketched in Fig. S1C).

Earlier version: A practical challenge of working with relative fitness is that it must be measured between all pairs of genotypes in coculture competitions. Therefore it is common to infer relative fitness of two genotypes based on some individual properties of the genotypes [70–73]. We denote this third notion of fitness as the fitness potential. Like absolute fitness, fitness potential is a property of an individual genotype, but unlike the absolute fitness, it has no meaning by itself. In analogy with the potential energy of physical systems, fitness potential is intended only to be compared between two genotypes (e.g., by taking the difference or ratio) to estimate the relative fitness between those genotypes [36, 37]. The collection of fitness potential values across a large set of genotypes forms a fitness landscape [74, 75]. We note that fitness as defined here gives information about short-time dynamics but not necessarily the long-term outcome (compare Smith et al. [6]). For example, this excludes the ratio of growth rates [7, 73] or the resource concentration R∗ in chemostat equilibrium [76] since these quantities cannot tell you how fast the absolute or relative abundance is changing. We also note that in all these definitions, the fitness of a genotype is specific to one environment, including all abiotic and biotic factors. How fitness in one environment is related to fitness in other environments is important [14, 19, 42, 43, 49, 50, 73, 77, 78] but beyond the scope of this article, as we focus on specific laboratory measurements of fitness rather than abstract concepts.

Revised version: A practical challenge of working with relative fitness is that it must be measured between all pairs of genotypes in coculture competitions. Therefore it is common to infer relative fitness of two genotypes based on some individual properties of the genotypes [70–73]. We denote this third notion of fitness as the fitness potential. Like absolute fitness, fitness potential is a property of an individual genotype, but unlike the absolute fitness, it has no meaning by itself. In analogy with the potential energy of physical systems, fitness potential is intended only to be compared between two genotypes (e.g., by taking the difference or ratio) to estimate the relative fitness between those genotypes [36, 37]. The collection of fitness potential values across a large set of genotypes forms a fitness landscape [74, 75]. Note that the term “fitness potential” has been used with a related but somewhat different meaning in the population genetics literature [76–78].

Added in the revised version: We emphasize that fitness as defined here gives information about short-time dynamics but not necessarily the long-term outcome (compare Smith et al. [6]). For example, this excludes the ratio of growth rates [7, 73] or the resource concentration R ∗ in chemostat equilibrium [79] since these quantities cannot tell you how fast the absolute or relative abundance is changing. We also note that in all these definitions, the fitness of a genotype is specific to one environment, including all abiotic and biotic factors. How fitness in one environment is related to fitness in other environments is important [14, 19, 42, 43, 49, 50, 73, 80, 81] but beyond the scope of this article, as we focus on specific laboratory measurements of fitness rather than abstract concepts.

Earlier version: In this article we will focus on relative fitness as this is the most common object of high-throughput laboratory experiments in microbes [18, 26, 40, 55, 56] and also has the most straightforward connection to key evolutionary outcomes such as fixation [57, 66, 67]. We define a genotype’s relative fitness as any number that is sufficient to predict the genotype’s relative abundance x(t) over a short time period (Figure S1). The simplest approach is to take a linear expansion of x(t) and use its slope at time t to predict the change over a time Δt (e.g., Figure 1A, top panel):

Revised version: In this article we will focus on relative fitness as this is the most common object of high-throughput laboratory experiments in microbes [18, 26, 40, 55, 56] and also has the most straightforward connection to key evolutionary outcomes such as fixation [57, 66, 67]. We define a genotype’s relative fitness as any number that is sufficient to predict the genotype’s relative abundance x(t) over a short time period (Fig. S1). The simplest approach is to take a linear expansion of x(t) and use its slope at time t to predict the change over a time Δt (e.g., Fig. 1A, top panel):

Removed in the revised version: (1) x(t+Δt)=x(t)+dxdt|tΔt The slope (2) slinear=dxdt|t qualifies as a fitness statistic but often will incorrectly estimate the actual change in relative abundance because x(t) is nonlinear (Figure 1A, top panel).

Overview of the choice of encoding and the choice of time scale for quantifying relative fitness.

(A) Example trajectory of relative abundance x (top panel) for a mutant invading and eventually replacing a wild-type population. The same trajectory is plotted under the encoding log x (middle panel) and the encoding logit x = log(x/(1 − x)) (bottom panel). (B) Flowchart to predict the future relative abundance of a mutant given a relative fitness value sm = dm/dt for some encoding m. The current relative abundance xt is transformed into the new variable mt = m(xt), then projected into the future through a linear extrapolation using sm (upper horizontal arrow) and finally converted back into a frequency xt+!t using the decoding function m−1. (C) Four scenarios for a mutant with the same relative fitness per-cycle but with different underlying population dynamics. For each scenario, we show an example trajectory of absolute abundance (stacked) for the wild-type (dark grey) and mutant population (light grey). Each scenario is mapped as a single-dot onto the fold-change diagram (center plot) and colored areas indicated positive (green area) and negative relative fitness per-cycle (blue area; compare Eq. (9)). Fitness per-cycle has isoclines that are parallel to the identity in the LFC diagram. (D) Scenario for rank discrepancy between relative fitness per-cycle scyclelogit and relative fitness per-generation sgenlogit. For a given competition (red dot), rank discrepancy occurs in a bow tie-shaped area of the fold-change space (shaded red). Any competition in the right half of this area (blue dot) will have higher mutant fitness scyclelogit but lower mutant fitness sgenlogit (right inset). As small plots on the left, we show possible population dynamics that generate this fold-change variation. Here the growth rates of the mutant and wild-type are identical in the two scenarios, but the growth period is different (e.g., from additional resources).

Added in the revised version: (1) x(t+Δt)=x(t)+dxdt|tΔt. The slope (2) slinear=dxdt|t qualifies as a fitness statistic but often will incorrectly estimate the actual change in relative abundance because x(t) is nonlinear (Fig. 1A, top panel).

Earlier version: Alternatively, we can encode the relative abundance x in a new variable m(x) (any smooth, strictly-increasing function) whose dynamics may be easier to linearly approximate (Figure 1B; compare Denniston and Crow [79]): (3) m(x(t+Δt))=m(x(t))+smΔt, where (4) sm=dmdt|t is the relative fitness under the encoding m. In Sec. S1 we show examples of the encoded relative fitness for an explicit model of population dynamics, and in Sec. S2 we show how this concept generalizes to absolute fitness as well.

Revised version: Alternatively, we can encode the relative abundance x in a new variable m(x) (any smooth, strictly-increasing function) whose dynamics may be easier to linearly approximate (Fig. 1B; compare Denniston and Crow [82]): (3) m(x(t+Δt))=m(x(t))+smΔt, where (4) sm=dmdt|t is the relative fitness under the encoding m. In Sec. S1 we show examples of the encoded relative fitness for an explicit model of population dyanmics, and in Sec. S2 we show how this concept generalizes to absolute fitness as well.

Earlier version: The ideal encoding transforms the relative abundance x(t) into a linear trajectory that can be fully predicted from the slope sm. While the log-encoding m(x) = log x is common in many TnSeq studies [17–19, 43, 77, 78] (and is implicit in plots of relative abundance on a logarithmic scale [44]), the log-encoded fitness slog tends to overestimate the change at high relative abundance (Figure 1A, middle panel). Instead, the best encoding in most cases is the logit function (Figure 1A, bottom panel): (5) logitx=log⁡(x1−x). This is because the relative abundance x(t) typically follows logistic dynamics (Sec. S1), at least as a first-order approximation (Sec. S3, Figure S2), and the logit function is linearly related to the inverse function of logistic x(t). The logit-encoded relative fitness of a mutant is equivalent to the standard definition of its selection coefficient in population genetics [80]. For this reason we focus on the logit-encoded relative fitness for the remainder of this work. Note that the logit encoding and logistic dynamics apply to relative abundance and are not to be confused with the logistic growth model used with microbial growth curves, where it describes the absolute abundance [81].

Revised version: The ideal encoding transforms the relative abundance x(t) into a linear trajectory that can be fully predicted from the slope sm. While the log-encoding m(x) = log x is common in many TnSeq studies [17–19, 43, 80, 81] (and is implicit in plots of relative abundance on a logarithmic scale [44]), the log-encoded fitness slog tends to overestimate the change at high relative abundance (Fig. 1A, middle panel). Instead, the best encoding in most cases is the logit function (Fig. 1A, bottom panel): (5) logitx=log⁡(x1−x) This is because the relative abundance x(t) typically follows logistic dynamics (Sec. S1), at least as a first-order approximation (Sec. S3, Fig. S2), and the logit function is linearly related to the inverse function of logistic x(t). The logit-encoded relative fitness of a mutant is equivalent to the standard definition of its selection coefficient in population genetics [83]. For this reason we focus on the logit-encoded relative fitness for the remainder of this work. Note that the logit encoding and logistic dynamics apply to relative abundance and are not to be confused with the logistic growth model used with microbial growth curves, where it describes the absolute abundance [84].

Estimating relative fitness requires a choice of time scale

Earlier version: Since the relative abundance of a genotype can only be measured at discrete times, we must estimate the instantaneous relative fitness of Eq. (4) using a finite difference between consecutive time points (Figure S1C, Sec. S4): (6) sm≈m(x(t))−m(x(t−Δt)Δt. For the logit-encoding this becomes (7) slogit=1Δt[log⁡x(t)1−x(t)−log⁡x(t−Δt)1−x(t−Δt)]. In the case of a single mutant competing against a wild-type strain, we can rewrite Eq. (7) in terms of the log fold-changes of the mutant and wild-type strains: (8) slogit=LFCmut−LFCwtΔt, where the log fold-change of each strain is LFC = log N (t)/N (t − Δt) and N is the biomass of the strain. The LFC is sometimes also referred to as a Malthusian parameter [9, 26, 65].

Revised version: Since the relative abundance of a genotype can only be measured at discrete times, we must estimate the instantaneous relative fitness of Eq. (4) using a finite difference between consecutive time points (Fig. S1C, Sec. S4): (6) sm≈m(x(t))−m(x(t−Δt)Δt. For the logit-encoding this becomes (7) slogit=1Δt[log⁡x(t)1−x(t)−log⁡x(t−Δt)1−x(t−Δt)]. In the case of a single mutant competing against a wild-type strain, we can rewrite Eq. (7) in terms of the log fold-changes of the mutant and wild-type strains: (8) slogit=LFCmut−LFCwtΔt, where the log fold-change of each strain is LFC = log N (t)/N (t − Δt) and N is the biomass of the strain. The LFC is sometimes also referred to as a Malthusian parameter [9, 26, 65].

Earlier version: Therefore a key element of estimating relative fitness from discrete time data is the choice of time interval Δt. There are two common choices in empirical measurements of microbial fitness. The first is to use a time scale extrinsic to the population. In particular, microbial populations in laboratory experiments [26, 82] but also in some natural environments [83–85] grow in discrete growth cycles dictated by external pulses of nutrients (e.g., batch culture). In this case, it is convenient to quantify relative fitness per-cycle, by choosing time units in Eq. (8) such that Δt = 1: (9) scyclelogit=LFCmut−LFCwt. The second common choice for Δt is a time scale intrinsic to the population, typically its generation time. The relative fitness per-cycle may depend on the number of generations in the growth cycle, and it can be valuable to normalize for this dependence, especially when comparing across environments [26]. However, defining the number of generations for a population with multiple genotypes growing at different rates is ambiguous. In the case of a single mutant competing with a wild-type, it is common to consider only the number of generations experienced by the wild-type strain, which is estimated by its log fold-change LFCwt. It is convenient to express the generations in base e to match the natural logarithm in the logit-encoding, but this could equivalently be done by converting all logarithms in Eq. (8) to base 2. Thus the relative fitness of a mutant per-generation is defined by choosing Δt = LFCwt in Eq. (8): (10) sgenlogit=LFCmutLFCwt−1. This quantity is equivalent to the fitness statistic W=1+sgenlogit used in the LTEE [26] and other studies [86–90]. Some authors assume a fixed number of generations per growth cycle [12, 15], but since this may not be true, we use the term “per-generation” strictly for fitness statistics where the wild-type generations are explicitly measured.

Revised version: Therefore a key element of estimating relative fitness from discrete time data is the choice of time interval Δt. There are two common choices in empirical measurements of microbial fitness. The first is to use a time scale extrinsic to the population. In particular, microbial populations in laboratory experiments [26, 85] but also in some natural environments [86−88] grow in discrete growth cycles dictated by external pulses of nutrients (e.g., batch culture). In this case, it is convenient to quantify relative fitness per-cycle, by choosing time units in Eq. (8) such that Δt = 1: (9) scyclelogit=LFCmut−LFCwt. The second common choice for Δt is a time scale intrinsic to the population, typically its generation time. The relative fitness per-cycle may depend on the number of generations in the growth cycle, and it can be valuable to normalize for this dependence, especially when comparing across environments [26]. However, defining the number of generations for a population with multiple genotypes growing at different rates is ambiguous. In the case of a single mutant competing with a wild-type, it is common to consider only the number of generations experienced by the wild-type strain, which is estimated by its log fold-change LFCwt. It is convenient to express the generations in base e to match the natural logarithm in the logit-encoding, but this could equivalently be done by converting all logarithms in Eq. (8) to base 2. Thus the relative fitness of a mutant per-generation is defined by choosing Δt = LFCwt in Eq. (8): (10) sgenlogit=LFCmutLFCwt−1. This quantity is equivalent to the fitness statistic W=1+sgenlogit used in the LTEE [26] and other studies [89–93]. Some authors assume a fixed number of generations per growth cycle [12, 15], but since this may not be true, we use the term “per-generation” strictly for fitness statistics where the wild-type generations are explicitly measured.

Relative fitness per-cycle and per-generation can rank mutants differently when wild-type growth depends on mutant

Earlier version: Although relative fitness of a mutant per growth-cycle (Eq. (9)) and per-generation (Eq. (10)) differ only by choice of time scale, they are not just equivalent quantities in different units. First, while fitness per-cycle is well-defined regardless of whether the population undergoes net growth or death (Figure 1C), fitness per-generation is not defined when the wild-type undergoes no change or net death (denominator of Eq. (10) is not positive) [5, 91]. This can occur for microbes under high drug concentrations, as well as for microbial populations under harsh environmental conditions as found in sediments or wastewater [92, 93].

Revised version: Although relative fitness of a mutant per growth-cycle (Eq. (9)) and per-generation (Eq. (10)) differ only by choice of time scale, they are not just equivalent quantities in different units. First, while fitness per-cycle is well-defined regardless of whether the population undergoes net growth or death (Fig. 1C), fitness per-generation is not defined when the wild-type undergoes no change or net death (denominator of Eq. (10) is not positive) [5, 94]. This can occur for microbes under high drug concentrations, as well as for microbial populations under harsh environmental conditions as found in sediments or wastewater [95, 96].

Earlier version: However, even when the wild-type LFC is positive, the relative fitness per-cycle and per-generation are not necessarily proportional to each other across a set of mutants competed against the same wild-type. The reason is that while the growth cycle time scale is extrinsic to a population and therefore the same across all pairs of genotypes, the generation time scale depends on the growth of the wild-type, which may be affected by the specific mutant it is competing against. This discrepancy can be quantitative as well as qualitative, producing different rankings of mutants. For example, Figure 1D shows two mutant genotypes (red and blue) that have opposite rankings under these fitness statistics: the blue mutant has higher relative fitness per-cycle scyclelogit, but the red mutant has higher relative fitness per-generation sgenlogit. The disagreement in ranking requires positive covariation in the LFCs (red bow tie-shaped area in Figure 1D), such that the mutant with higher LFC also induces a higher LFC in the wildtype (see illustration in Figure and exact conditions in Sec. S5.

Revised version: However, even when the wild-type LFC is positive, the relative fitness per-cycle and per-generation are not necessarily proportional to each other across a set of mutants competed against the same wild-type. The reason is that while the growth cycle time scale is extrinsic to a population and therefore the same across all pairs of genotypes, the generation time scale depends on the growth of the wild-type, which may be affected by the specific mutant it is competing against. This discrepancy can be quantitative as well as qualitative, producing different rankings of mutants. For example, Fig. 1D shows two mutant genotypes (red and blue) that have opposite rankings under these fitness statistics: the blue mutant has higher relative fitness per-cycle scyclelogit, but the red mutant has higher relative fitness per-generation sgenlogit. The disagreement in ranking requires positive covariation in the LFCs (red bow tie-shaped area in Fig. 1D), such that the mutant with higher LFC also induces a higher LFC in the wild-type (see illustration in Fig. and exact conditions in Sec.S5)

Basic features of microbial population dynamics are sufficient to cause rank discrepancy between fitness statistics

Since the discrepancy between fitness per-cycle scyclelogit and per-generation sgenlogit occurs when the wild-type LFC varies across pairwise competitions with different mutants, we need to determine what biological mechanisms cause this. One way to explore this is to use a minimal model of microbial population dynamics parameterized by empirical data on a set of mutants. While a minimal model does not capture all possible mechanisms, it provides a null expectation for how common these discrepancies should be in real populations.

Earlier version: We use a model that approximates microbial population dynamics as three growth phases — a lag phase, an exponential growth phase, and a stationary phase (Figure 2A) — where strains compete for a single limiting resource with no other interactions [94, 95], since resource competition is likely to be the main interaction between a wild-type and a mutant with only small genetic differences (e.g., a point mutation or gene knockout). This model captures the minimal dynamics of microbial growth under laboratory conditions, and therefore serves as a baseline for possible outcomes. We parameterize the model using growth curves of single-gene knockout mutants in Saccharomyces cerevisiae [29, 96] (Figures and S5, Sec. S6). We simulate a pairwise competition for each knockout against the wild-type (Figure 2A) and calculate relative fitness per-cycle scyclelogit and per-generation sgenlogit for each knockout (Methods). While the relative fitness per-cycle and per-generation in the model are highly correlated across mutants overall (Figure S6), the model predicts major differences in individual mutants when we rank them by fitness (Figure 2B). Indeed, when we plot all simulated mutant competitions on the LFC diagram, we find that there are many pairs of points that have positive covariation in the mutant and wild-type LFC and thus give rise to rank differences (compare Figures 2C and 1D). Since these rankings often serve to select important mutants from a functional genomics screen or an evolution experiment, discrepancies between fitness definitions may lead analyses down different paths.

Revised version: We use a model that approximates microbial population dynamics as three growth phases — a lag phase, an exponential growth phase, and a stationary phase (Fig. 2A) — where strains compete for a single limiting resource with no other interactions [97, 98], since resource competition is likely to be the main interaction between a wild-type and a mutant with only small genetic differences (e.g., a point mutation or gene knockout). This model captures the minimal dynamics of microbial growth under laboratory conditions, and therefore serves as a baseline for possible outcomes. We parameterize the model using growth curves of single-gene knockout mutants in Saccharomyces cerevisiae [29, 99] (Figs. and S5, Sec. S6). We simulate a pairwise competition for each knockout against the wild-type (Fig. 2A) and calculate relative fitness per-cycle scyclelogit and per-generation sgenlogit for each knockout (Methods). While the relative fitness per-cycle and per-generation in the model are highly correlated across mutants overall (Fig. S6), the model predicts major differences in individual mutants when we rank them by fitness (Fig. 2B). Indeed, when we plot all simulated mutant competitions on the LFC diagram, we find that there are many pairs of points that have positive covariation in the mutant and wild-type LFC and thus give rise to rank differences (compare Figs. 2C and 1D). Since these rankings often serve to select important mutants from a functional genomics screen or an evolution experiment, discrepancies between fitness definitions may lead analyses down different paths.

Comparison of mutant fitness rankings with different time scales using empirical trait variation.

Earlier version: (A) Overview of pairwise coculture simulations. For each mutant strain (grey), we simulate a competition growth cycle against a reference wild-type strain (blue) using the estimated traits and laboratory parameters for the initial condition (N0 = 0.05 OD, R0 = 111 mM glucose, x = 0.5; Methods and see Figure for the underlying trait distribution). (B) Rank discrepancy between relative fitness per-cycle scyclelogit and per-generation sgenlogit. Higher rank is defined as higher fitness, and the rank difference is defined as the rank in sgenlogit minus the rank in scyclelogit. The blue halo marks the mutant with the greatest rank difference. (C) Covariation between wild-type and mutant fold-change across all simulated competitions for all mutant strains (grey dots). The blue halo marks the same mutant as in panel B, with its red bow tie-shaped area defining the space of LFCs that would have a rank discrepancy with it (compare Figure 1D), which includes many other mutant competitions.

Revised version: (A) Overview of pairwise coculture simulations. For each mutant strain (grey), we simulate a competition growth cycle against a reference wild-type strain (blue) using the estimated traits and laboratory parameters for the initial condition (N0 = 0.05 OD, R0 = 111 mM glucose, x = 0.5; Methods and see Fig. for the underlying trait distribution). (B) Rank discrepancy between relative fitness per-cycle scyclelogit and per-generation sgenlogit. Higher rank is defined as higher fitness, and the rank difference is defined as the rank in sgenlogit minus the rank in scyclelogit. The blue halo marks the mutant with the greatest rank difference. (C) Covariation between wild-type and mutant fold-change across all simulated competitions for all mutant strains (grey dots). The blue halo marks the same mutant as in panel B, with its red bow tie-shaped area defining the space of LFCs that would have a rank discrepancy with it (compare Fig. 1D), which includes many other mutant competitions.

Earlier version: This model therefore shows that the basic features of microbial population dynamics and resource competition with realistic parameters are sufficient to cause substantial disagreement between fitness per-cycle and per-generation, underscoring the importance of this analysis choice. Furthermore, the model reveals a basic mechanism that causes this discrepancy: mutants must have diminished biomass yield plus variation in one additional trait (lag time or growth rate), and a high mutant relative abundance (e.g., 50%) at the start of the competition (Figure S7). These conditions are necessary because they allow mutants both to have variation in fitness and to alter the wild-type LFC in pairwise competition. Even though the wild-type growth rate is unaffected by mutants, the wild-type LFC will change in the presence of mutants that have lower yield because the population on average consumes more resources per cell division (Sec. S7). By this mechanism, it is possible to concoct a set of growth traits where relative fitness per-generation and per-cycle deliver completely opposite rankings (Figure S8).

Revised version: This model therefore shows that the basic features of microbial population dynamics and resource competition with realistic parameters are sufficient to cause substantial disagreement between fitness per-cycle and pergeneration, underscoring the importance of this analysis choice. Furthermore, the model reveals a basic mechanism that causes this discrepancy: mutants must have diminished biomass yield plus variation in one additional trait (lag time or growth rate), and a high mutant relative abundance (e.g., 50%) at the start of the competition (Fig. S7). These conditions are necessary because they allow mutants both to have variation in fitness and to alter the wild-type LFC in pairwise competition. Even though the wild-type growth rate is unaffected by mutants, the wild-type LFC will change in the presence of mutants that have lower yield because the population on average consumes more resources per cell division (Sec. S7). By this mechanism, it is possible to concoct a set of growth traits where relative fitness per-generation and per-cycle deliver completely opposite rankings (Fig. S8).

Fitness statistics can disagree on epistasis and long-term fitness trends

Earlier version: Since the relative fitness of single mutants differs depending on whether fitness is quantified per-cycle or pergeneration, how does that choice impact the apparent epistatic interactions between multiple mutations and the long-term evolutionary trend in fitness? Using the same minimal model of population dynamics as in Figure 2A, we simulate a set of six mutants, each one having an increase or decrease in each of the three quantitative traits in the model (lag time, growth rate, and yield), and calculate the relative fitness per-cycle and per-generation for all pairs of those mutations (assuming the mutation effects on the traits themselves are additive; Methods). We quantify epistasis as the fitness difference (in 1:1 competition with the wild-type; Methods) between the double mutant and the sum of the single mutants (Figure 3A) [97].

Revised version: Since the relative fitness of single mutants differs depending on whether fitness is quantified percycle or per-generation, how does that choice impact the apparent epistatic interactions between multiple mutations and the long-term evolutionary trend in fitness? Using the same minimal model of population dynamics as in Fig. 2A, we simulate a set of six mutants, each one having an increase or decrease in each of the three quantitative traits in the model (lag time, growth rate, and yield), and calculate the relative fitness per-cycle and per-generation for all pairs of those mutations (assuming the mutation effects on the traits themselves are additive; Methods). We quantify epistasis as the fitness difference (in 1:1 competition with the wild-type; Methods) between the double mutant and the sum of the single mutants (Fig. 3A) [100].

Discrepancy between fitness per-generation and per-cycle in epistasis detection and long-term trend.

Earlier version: (A) Schematic detection of epistasis for two single mutants (light and dark circle). Epistasis occurs if the corresponding double mutant (two-component circles ) deviates from the additive fitness (parallelogram shape in solid lines; Methods). (B) Correlation between epistasis in relative fitness per-cycle scycle and per-generation sgen. Each dot corresponds to a hypothetical double mutant that combines pairs of mutations systematically perturbing each of the three growth traits (Figure 2A) in our minimal model of population dynamics (Methods). (C) Heatmap for epistasis in the per-cycle fitness scycle for all double mutants in our simulated data set, same data as in panel B. (D) Heatmap for epistasis in the per-generation fitness sgen, same data as in panel B. All epistasis plots are based on 1:1 competition growth cycles with the wild-type (Figures S9–S14). See Figure for epistasis patterns at low mutant relative abundance. (E) Hypothetical scenario for trait evolution in an evolution experiment. An evolving population decreases lag time to zero (orange line), increases growth rate until saturation (blue line), and gradually decreases biomass yield (green line). This trend is similar to initial observations from the LTEE [98]. (F) Corresponding long-term trend in relative fitness based on the trait evolution in panel E. We estimate relative fitness per-cycle scyclelogit (grey line) and per-generation sgenlogit (red line) every 250 generations. Dotted vertical lines mark the end of trait evolution in lag time and growth rate in panel E. For the actual fitness trend in the LTEE, see Figure S16.

Revised version: (A) Schematic detection of epistasis for two single mutants (light and dark circle). Epistasis occurs if the corresponding double mutant (two-component circles) deviates from the additive fitness (parallelogram shape in solid lines; Methods). (B) Correlation between epistasis in relative fitness per-cycle scycle and per-generation sgen. Each dot corresponds to a hypothetical double mutant that combines pairs of mutations systematically perturbing each of the three growth traits (Fig. 2A) in our minimal model of population dynamics (Methods). (C) Heatmap for epistasis in the per-cycle fitness scycle for all double mutants in our simulated data set, same data as in panel B. The rows and columns correspond to the six single mutations in our model of population dynamics (Methods): an increase or decrease in lag time (+ω, − ω), an increase or decrease in maximum growth rate (+g, − g), or an increase or decrease in biomass yield (+Y, −Y). (D) Heatmap for epistasis in the per-generation fitness sgen, same data as in panel B. All epistasis plots are based on 1:1 competition growth cycles with the wild-type (Figs. –). See Fig. for epistasis patterns at low mutant relative abundance. (E) Hypothetical scenario for trait evolution in an evolution experiment. An evolving population decreases lag time to zero (orange line), increases growth rate until saturation (blue line), and gradually decreases biomass yield (green line). This trend is similar to initial observations from the LTEE [101]. (F) Corresponding long-term trend in relative fitness based on the trait evolution in panel E. We estimate relative fitness per-cycle Scycle (grey line) and per-generation sgenlogit (red line) every 250 generations. Dotted vertical lines mark the end of trait evolution in lag time and growth rate in panel E. For the actual fitness trend in the LTEE, see Fig.

Earlier version: Figure 3B shows that almost all the double mutants differ in their apparent epistasis per-cycle vs. per-generation. Epistasis per-cycle occurs when one of the mutations affects growth rate (Figure 3C), while epistasis per-generation occurs when one of the mutations affects lag time (Figure 3D). This is because the effect of a change in growth rate accrues per-generation (and hence is non-epistatic when measured over that time scale; Sec. S8), while the effect of a change in lag time accrues once per-cycle (and hence is non-epistatic in the per-cycle statistic). See Figures S9–S14 for growth curves in all cases. As an example, the double mutant in Figure S12, row D has two lag time-reducing mutations that are perfectly additive in the per-cycle statistic. However, the lag time mutations also reduce the saturation time of the growth cycle such that the fitness of the double mutant is spread over a fewer number of wild-type generations LFCwt, creating super-additive epistasis in the per-generation statistic. Note that this depends on measuring fitness for a mutant at relatively high abundance in the pairwise competition with the wild-type; when the mutant is rare, the wild-type fold-change LFCwt is effectively constant across mutants and epistasis measured per-cycle is proportional to its value measured per-generation (Figure S15). Differences in apparent epistasis can also lead to differences in apparent long-term trends of fitness per-cycle and per-generation in an evolution experiment. For example, if growth traits of a population in the minimal model evolve according to Figure 3E, then fitness per-cycle will saturate at an early time while fitness per-generation will continue to increase (Figure 3F). This is because the long-term decrease in yield reduces the wild-type LFC in 1:1 competitions between an evolved mutant and the wild-type, which causes fitness per-generation to increase (Eq. (10)).

Revised version: Figure 3B shows that almost all the double mutants differ in their apparent epistasis per-cycle vs. per-generation. Epistasis per-cycle occurs when one of the mutations affects growth rate (Fig. 3C), while epistasis pergeneration occurs when one of the mutations affects lag time (Fig. 3D). This is because the effect of a change in growth rate accrues per-generation (and hence is non-epistatic when measured over that time scale; Sec. S8), while the effect of a change in lag time accrues once per-cycle (and hence is non-epistatic in the percycle statistic). See Figs. S9–S14 for growth curves in all cases. As an example, the double mutant in Fig. S12, row D has two lag time-reducing mutations that are perfectly additive in the per-cycle statistic. However, the lag time mutations also reduce the saturation time of the growth cycle such that the fitness of the double mutant is spread over a fewer number of wild-type generations LFCwt, creating super-additive epistasis in the per-generation statistic. Note that this depends on measuring fitness for a mutant at relatively high abundance in the pairwise competition with the wild-type; when the mutant is rare, the wild-type fold-change LFCwt is effectively constant across mutants and epistasis measured per-cycle is proportional to its value measured per-generation (Fig. S15). Differences in apparent epistasis can also lead to differences in apparent long-term trends of fitness per-cycle and per-generation in an evolution experiment. For example, if growth traits of a population in the minimal model evolve according to Fig. 3E, then fitness per-cycle will saturate at an early time while fitness per-generation will continue to increase (Fig. 3F). This is because the long-term decrease in yield reduces the wild-type LFC in 1:1 competitions between an evolved mutant and the wild-type, which causes fitness per-generation to increase (Eq. (10)).

Different fitness time scales agree on LTEE trend but not on mutant ranking

Earlier version: As an empirical test for discrepancies between relative fitness per-generation and per-cycle, we test how the choice of fitness time scale affects conclusions from the LTEE (Sec. S9). We reanalyze competition data from the LTEE [11] using the relative fitness per-generation sgenlogit (equivalent to the original definition used in those studies, W=1+sgenlog⁡it) and relative fitness per-cycle sgenlogit. Note these competitions are between the ancestral wild-type and each of the 12 evolved populations, rather than single mutants, across time points. We confirm that the long-term fitness trend in the LTEE is the same for scyclelogit (Figure S16). When we look more closely at the fitness ranking between lineages, we find that the two statistics disagree on the fitness ranking at almost every time point (Figure 4A) and the fittest population overall (Figure 4B; see Figure S17A for rankings). This discrepancy arises from positive covariation across pairs of LFC values in the LTEE competition data set (Figure S17B).

Revised version: As an empirical test for discrepancies between relative fitness per-generation and per-cycle, we test how the choice of fitness time scale affects conclusions from the LTEE (Sec. S9). We reanalyze competition data from the LTEE [11] using the relative fitness per-generation sgenlogit (equivalent to the original definition used in those studies, W=1+sgenlogit) and relative fitness per-cycle scyclelogit. Note these competitions are between the ancestral wild-type and each of the 12 evolved populations, rather than single mutants, across time points. We confirm that the long-term fitness trend in the LTEE is the same for scyclelogit (Fig. S16). When we look more closely at thefitness ranking between lineages, we find that the two statistics disagree on the fitness ranking at almost every time point (Fig. 4A) and the fittest population overall (Fig. 4B; see Fig. S17A for rankings). This discrepancy arises from positive covariation across pairs of LFC values in the LTEE competition data set (Fig. S17B).

Rank discrepancies between fitness statistics in LTEE competition data.

Earlier version: (A) Rank difference between relative fitness per-cycle and per-generation as a function of time in the LTEE. For each of the 12 lines in the LTEE, we compute the relative fitness per-cycle scyclelogit and per-generation sgenlogit as a function of time, by averaging the data from the LTEE competition data set [11] across replicate measurements (Sec. S9). For some lines, the time-series is truncated due to measurement difficulties [11]. Based on the quantitative fitness values scyclelogit and sgenlogit, we compute fitness rankings between the replicate lines at any given time point (higher rank means higher fitness ; compare Figure 2B). The rank difference is defined as the rank in sgenlogit minus the rank in scyclelogit, and we combine the rank difference between all pairs of evolving lines to compute the maximum rank difference at each time-point. Insets: At three chosen time-points (t = 4000, t = 15000, t = 30000 generations), we show the correlation between relative fitness per-cycle scyclelogit and relative fitness per-generation sgenlogit (each dot corresponds to one of the 12 lines in the LTEE). We quantify the correlation between scyclelogit and sgenlogit with the Spearman rank correlation λ (panel title). Colors indicate the time-point and correspond to the color bar in panel B. (B) Covariation between relative fitness per-cycle scyclelogit and per-generation sgenlogit for evolved populations of the LTEE. Each line of the LTEE contributes roughly two competitions per time-point due to replicate measurements (Sec. S9) [11]. The color of the points indicates evolutionary time point; see panel (C) for the color bar. We highlight the evolved population with the greatest rank difference (blue halo; same as in Figure S17). (C) Same as panel B but comparing relative fitness under the logit-encoding scyclelogit and under the log-encoding scyclelog.

Revised version: (A) Rank difference between relative fitness per-cycle and per-generation as a function of time in the LTEE. For each of the 12 lines in the LTEE, we compute the relative fitness per-cycle Scycle and per-generation sgenlogit as a function of time, by averaging the data from the LTEE competition data set [11] across replicate measurements (Sec. S9). For some lines, the time-series is truncated due to measurement difficulties [11]. Based on the quantitative fitness values Scycle and sgenlogit, we compute fitness rankings between the replicate lines at any given time point (higher rank means higher fitness; compare Fig. 2B). The rank difference is defined as the rank in sgenlogit minus the rank in scyclelogit, and we combine the rank difference between all pairs of evolving lines to compute the maximum rank difference at each time-point. Insets: At three chosen time-points (t = 4000, t = 15000, t = 30000 generations), we show the correlation between relative fitness per-cycle Scycle and relative fitness per-generation sgenlogit (each dot corresponds to one of the 12 lines in the LTEE). We quantify the correlation between scyclelogit and sgenlogit with the Spearman rank correlation ε (panel title). Colors indicate the time-point and correspond to the color bar in panel B. (B) Covariation between relative fitness per-cycle scyclelogit and per-generation sgenlogit for evolved populations of the LTEE. Each line of the LTEE contributes roughly two competitions per time-point due to replicate measurements (Sec. S9) [11]. The color of the points indicates evolutionary time point; see panel (C) for the color bar. We highlight the evolved population with the greatest rank difference (blue halo; same as in Fig. S17). (C) Same as panel B but comparing relative fitness under the logit-encoding scyclelogit and under the log-encoding scyclelog.

Earlier version: While we have focused on the logit encoding so far since logistic dynamics are usually the best null model of mutant frequency dynamics (Figure 1A), the log encoding is common in many TnSeq studies [17–19, 43, 77, 78]. The LTEE competition data is a valuable opportunity to empirically test the consequences of this choice as well. Figure 4C shows that the relative fitness per-cycle under the log-encoding scyclelog differs from the logit-based fitness scyclelogit for evolved LTEE populations, leading to ranking discrepancies especially for the fittest populations (Figure S18A). The discrepancy in fitness between encodings has a different mechanism compared to the discrepancy from fitness time scale: it arises from variation in the initial relative abundance in mutant-to-wild-type competitions in the LTEE data set. Although these pairwise competitions are intended to start with equal abundances of the evolved population and of the wild-type (1:1), Figure S18B shows that there is a systematic trend toward lower initial frequencies of the evolved populations over time. This trend is a side effect of the inoculation method in LTEE competition experiments, where a fixed volume of the evolved population is transferred to the competition growth cycle but the cell density for the evolved population in pre-culture decreases [98, 99].

Revised version: While we have focused on the logit encoding so far since logistic dynamics are usually the best null model of mutant frequency dynamics (Fig. 1A), the log encoding is common in many TnSeq studies [17–19, 43, 80, 81]. The LTEE competition data is a valuable opportunity to empirically test the consequences of this choice as well. Figure 4C shows that the relative fitness per-cycle under the log-encoding scyclelog differs from the logit-based fitness scyclelogit for evolved LTEE populations, leading to ranking discrepancies especially for the fittest populations (Fig. S18A). The discrepancy in fitness between encodings has a different mechanism compared to the discrepancy from fitness time scale: it arises from variation in the initial relative abundance in mutant-to-wild-type competitions in the LTEE data set. Although these pairwise competitions are intended to start with equal abundances of the evolved population and of the wild-type (1:1), Fig. S18B shows that there is a systematic trend toward lower initial frequencies of the evolved populations over time. This trend is a side effect of the inoculation method in LTEE competition experiments, where a fixed volume of the evolved population is transferred to the competition growth cycle but the cell density for the evolved population in pre-culture decreases [101, 102].

Fitness in bulk competitions is distorted by higher-order effects depending on reference subpopulation

Earlier version: So far we have focused on relative fitness in pairwise competition experiments, but this is often not practical for large numbers of mutants due to the required strain labeling (e.g., with fluorescent proteins or selection markers). One alternative is to measure a fitness potential (as defined at the beginning of the Results section), such as growth rate or area under the growth curve (AUC) [32, 34, 36, 37], from monoculture experiments, under the assumption that true relative fitness is proportional to the difference in this quantity between genotypes. Using our minimal model of population dynamics parameterized by single-gene knockouts in yeast (Figure 2A), we find that common fitness potentials perform poorly overall (Figure S19; see refs. [32, 34, 36, 37] for empirical tests), with some success for the AUC but strongly depending on the choice of time window (Figure S20, Sec. S10).

Revised version: So far we have focused on relative fitness in pairwise competition experiments, but this is often not practical for large numbers of mutants due to the required strain labeling (e.g., with fluorescent proteins or selection markers). One alternative is to measure a fitness potential (as defined at the beginning of the Results section), such as growth rate or area under the growth curve (AUC) [32, 34, 36, 37], from monoculture experiments, under the assumption that true relative fitness is proportional to the difference in this quantity between genotypes. Using our minimal model of population dynamics parameterized by single-gene knockouts in yeast (Fig. 2A), we find that common fitness potentials perform poorly overall (Fig. S19; see refs. [32, 34, 36, 37] for empirical tests), with some success for the AUC but strongly depending on the choice of time window (Fig. S20, Sec. S10).

Earlier version: Another approach to measuring relative fitness for a large number of mutants is to use bulk competition experiments, where many mutant genotypes compete simultaneously in a single culture and each genotype is tracked through high-throughput sequencing [39, 44]. This raises the question of how well relative fitness of a mutant in bulk competition corresponds to its relative fitness in a pairwise competition with the wild-type, since the growth in bulk might be influenced by the presence of other mutant genotypes, a phenomenon known as a higher-order interaction [78, 95, 100].

Revised version: Another approach to measuring relative fitness for a large number of mutants is to use bulk competition experiments, where many mutant genotypes compete simulta-neously in a single culture and each genotype is tracked through high-throughput sequencing [39, 44]. This raises the question of how well relative fitness of a mutant in bulk competition corresponds to its relative fitness in a pairwise competition with the wild-type, since the growth in bulk might be influenced by the presence of other mu-tant genotypes, a phenomenon known as a higher-order interaction [81, 98, 103].

Earlier version: An important choice in bulk competition experiments is the relative abundance of the library of all mutant genotypes compared to the wild-type. The invasion of a spontaneous mutation into an existing population is best captured by pairwise competitions with low initial relative abundance of the mutant (Case I, Figure 5A), and one way to recreate this scenario in a bulk competition is to use a low relative abundance of the mutant library overall (Case II, Figure 5A) [40, 48–52]. A practical problem with Case II is that individual genotypes in the library will have low absolute abundances, which leads to stochasticity in the population dynamics and the sequencing process [13, 101, 102]. Therefore, it is common to compete the mutant library by itself (Case III, Figure 5A) [13, 15, 17–19, 77].

Revised version: An important choice in bulk competition experiments is the relative abundance of the library of all mutant genotypes compared to the wild-type. The invasion of a spontaneous mutation into an existing population is best captured by pairwise competitions with low initial relative abundance of the mutant (Case I, Fig. 5A), and one way to recreate this scenario in a bulk competition is to use a low relative abundance of the mutant library overall (Case II, Fig. 5A) [40, 48–52]. A practical problem with Case II is that individual genotypes in the library will have low absolute abundances, which leads to stochasticity in the population dynamics and the sequencing process [13, 104, 105]. Therefore, it is common to compete the mutant library by itself (Case III, Fig. 5A) [13, 15, 17–19, 80].

The choice of library abundance and reference group in bulk competition experiments.

Earlier version: (A) Overview of a pairwise competition experiment with one mutant (top row) and scenarios for bulk competition experiments with many different mutants (middle and bottom rows). The two bulk competition scenarios differ in their initial fraction of the mutant library (colored ovals) in the inoculum (open box). For each scenario, we show a schematic growth cycle (log absolute abundance) in the inset on the right. (B) Schematic relative abundance trajectories for a mutant compared to two alternative subpopulations. We distinguish between the total relative abundance xi with respect to the population as a whole (height of green band in the top box) and the pairwise relative abundance xi,wt with respect to the wild-type (height of green band in the bottom box; Eq. (20)). We indicate the sign of total relative fitness (Eq. (24)) and pairwise relative fitness (Eq. (25)) on the right. (C) The absolute error between bulk and pairwise competition experiments. The total relative fitness (dark grey dots; Eq. (24)) and the pairwise relative fitness (orange dots; Eq. (25)) for mutants in the minimal population dynamics model (Figure 2A, Methods) parameterized by the yeast knockout data (Figure S4) in bulk competition growth cycle with low mutant library abundance (panel A, case II; Methods). The absolute error is defined as the bulk fitness statistic minus the relative fitness in pairwise competition (Eq. (S92) in Sec. S15). In the inset, the absolute error for pairwise relative fitness (Eq. (25)) for a bulk competition growth cycle with 99.9% library abundance (light grey dots; case III). This still includes a wild-type reference at small percentage to be able to compute a pairwise relative fitness. The x-axis and the orange dots in the inset are identical to the main plot. (D) The relative error in bulk competition experiments as a function of mutant library abundance in the inoculum. Each line corresponds to a knockout mutant in our data set, and represents the relative error between the pairwise relative fitness in bulk competition and the relative fitness in pairwise competition (Eq. (S102) in Sec. S16). The black vertical lines show the recommended mutant library abundance for our data set based on Eq. (11) (xlib ≈ 24.6%) and based on the more conservative Eq. (S118) (xlib ≈ 0.02%, Sec. S16). Note these bounds are calculated a priori and are not the result of a fitting procedure.

Revised version: (A) Overview of a pairwise competition experiment with one mutant (top row) and scenarios for bulk competition experiments with many different mutants (middle and bottom rows). The two bulk competition scenarios differ in their initial fraction of the mutant library (colored ovals) in the inoculum (open box). For each scenario, we show a schematic growth cycle (log absolute abundance) in the inset on the right. (B) Schematic relative abundance trajectories for a mutant compared to two alternative subpopulations. We distinguish between the total relative abundance xi with respect to the population as a whole (height of green band in the top box) and the pairwise relative abundance xi,wt with respect to the wild-type (height of green band in the bottom box; Eq. (20)). We indicate the sign of total relative fitness (Eq. (24)) and pairwise relative fitness (Eq. (25)) on the right. (C) The absolute error between bulk and pairwise competition experiments. The total relative fitness (dark grey dots; Eq. (24)) and the pairwise relative fitness (orange dots; Eq. (25)) for mutants in the minimal population dynamics model (Fig. 2A, Methods) parameterized by the yeast knockout data (Fig. S4) in bulk competition growth cycle with low mutant library abundance (panel A, case II; Methods). The absolute error is defined as the bulk fitness statistic minus the relative fitness in pairwise competition (Eq. (S92) in Sec. S15). In the inset, the absolute error for pairwise relative fitness (Eq. (25)) for a bulk competition growth cycle with 99.9% library abundance (light grey dots; case III). This still includes a wild-type reference at small percentage to be able to compute a pairwise relative fitness. The x-axis and the orange dots in the inset are identical to the main plot. (D) The relative error in bulk competition experiments as a function of mutant library abundance in the inoculum. Each line corresponds to a knockout mutant in our data set, and represents the relative error between the pairwise relative fitness in bulk competition and the relative fitness in pairwise competition (Eq. (S102) in Sec. S16). The black vertical lines show the recommended mutant library abundance for our data set based on Eq. (11) (xlib ≈ 24.6%) and based on the more conservative Eq. (S118) (xlib ≈ 0.02%, Sec. S16). Note these bounds are calculated a priori and are not the result of a fitting procedure.

Earlier version: A second important choice is whether to quantify fitness of a mutant relative to the whole population (“total relative fitness”) [18, 19, 77] or relative to another specific genotype, like the wild-type (“pairwise relative fitness”; see Figure 5B) [40, 48–52]. In the case of a mutant library growing by itself (Case III), it is still possible to estimate a pairwise fitness by using a few known neutral mutants as the reference population [15, 17] and calculating fitness relative to this group (Methods; Secs. and S12).

Revised version: A second important choice is whether to quantify fitness of a mutant relative to the whole population (“total relative fitness”) [18, 19, 80] or relative to another specific genotype, like the wild-type (“pairwise relative fitness”; see Fig. 5B) [40, 48–52]. In the case of a mutant library growing by itself (Case III), it is still possible to estimate a pairwise fitness by using a few known neutral mutants as the reference population [15, 17] and calculating fitness relative to this group (Methods; Secs. and S12).

Earlier version: To test the consequences of these choices, we use the minimal population dynamics model (Figure 2A) parameterized by yeast single-gene knockouts and compare the estimates of total and pairwise relative fitness in bulk competitions to fitness measured in pairwise competition, which we use as the ground truth (Methods). Figure 5C shows that in bulk competitions, the error in total relative fitness (compared to fitness in pairwise competitions) is systematically offset from the error in pairwise relative fitness. For example, a mutant that grows identically to the wild-type (neutral phenotype) does not change in abundance relative to the wild-type, but it may have a net increase in its total relative abundance due to the poor growth of other mutants (compare top and bottom panel in Figure 5B). This offset affects the total relative fitness of all mutants in a uniform way (Figure 5C), such that total and pairwise relative fitness completely agree in the ranking of genotypes (Figure S21). Note that this off-set between total and relative fitness is general and does not depend on our minimal population dynamics model (Sec. S13).

Revised version: To test the consequences of these choices, we use the minimal population dynamics model (Fig. 2A) parameterized by yeast single-gene knockouts and compare the estimates of total and pairwise relative fitness in bulk competitions to fitness measured in pairwise competition, which we use as the ground truth (Methods). Figure 5C shows that in bulk competitions, the error in total relative fitness (compared to fitness in pairwise competitions) is systematically offset from the error in pairwise relative fitness. For example, a mutant that grows identically to the wild-type (neutral phenotype) does not change in abundance relative to the wild-type, but it may have a net increase in its total relative abundance due to the poor growth of other mutants (compare top and bottom panel in Fig. 5B). This offset affects the total relative fitness of all mutants in a uniform way (Fig. 5C), such that total and pairwise relative fitness completely agree in the ranking of genotypes (Fig. S21). Note that this offset between total and relative fitness is general and does not depend on our minimal population dynamics model (Sec. S13).

Earlier version: While the pairwise relative fitness is the preferred method for estimating fitness in bulk competitions, the presence of higher-order interactions means it still deviates from the fitness in pairwise competitions. Figure 5C shows the absolute error for pairwise fitness has a positive trend with fitness and increases with the relative abundance of the mutant (Case III, inset in Figure 5C). For the specific population dynamics in our simulation, we decompose the higher-order interaction into two terms (Secs. and S15): a fitness-independent term and a fitness-dependent term that acts as an amplifier for fitness values (Figure S22) and drives the positive error trend in Figure 5C. Intuitively, the population in bulk competition consumes resources more slowly and gives more time for growth rate differences to accrue. The amplification of mutant pairwise fitness in bulk is important because it widens the apparent distribution of fitness effects (leading to, for example, higher estimates for the speed of adaptation) than if fitness was measured in pairwise competition.

Revised version: While the pairwise relative fitness is the preferred method for estimating fitness in bulk competitions, the presence of higher-order interactions means it still deviates from the fitness in pairwise competitions. Figure 5C shows the absolute error for pairwise fitness has a positive trend with fitness and increases with the relative abundance of the mutant (Case III, inset in Fig. 5C). For the specific population dynamics in our simulation, we decompose the higher-order interaction into two terms (Secs. and S15): a fitness-independent term and a fitness-dependent term that acts as an amplifier for fitness values (Fig. S22) and drives the positive error trend in Fig. 5C. Intuitively, the population in bulk competition consumes resources more slowly and gives more time for growth rate differences to accrue. The amplification of mutant pairwise fitness in bulk is important be-cause it widens the apparent distribution of fitness effects (leading to, for example, higher estimates for the speed of adaptation) than if fitness was measured in pairwise competition.

Earlier version: In order to control the error from higher-order inter-actions, we derive an explicit error bound. Figure 5D shows the relative error from higher-order interactions for each mutant across a range of mutant library abundances. Using the specific population dynamics of our model, we derive the following estimate: Assuming that the mutants in the library only have variation in growth rate, the mutant library abundance should satisfy (11) xlib<ϵ⋅gwt|glib−gwt|, where ε is the desired threshold on the relative error (Figure S23), gwt is the wild-type growth rate, and glib is the growth rate of the library as a whole (Sec. S16). In the case of our single-gene knockout library, Eq. (11) predicts a maximum library abundance of 24.6% based on the wild-type and library growth rate (gwt = 0.406, glib = 0.389) for a relative error ε = 1%. Figure 5D shows that this maximum library abundance keeps the relative error below 1% for high-fitness mutants (bright yellow), because they are dominated by growth rate effects, but fails for mutants close to neutrality because they have a tradeoff between growth rate and lag time, which the estimate in Eq. (11) neglects. It is possible to derive a more precise bound that keeps the relative error below 1% for all mutant genotypes (in this case xlib = 0.01%, see vertical line in Figure 5D), but this requires prior knowledge of the trait covariation in the mutant library (Sec. S16).

Revised version: In order to control the error from higher-order interactions, we derive an explicit error bound. Figure 5D shows the relative error from higher-order interactions for each mutant across a range of mutant library abundances. Using the specific population dynamics of our model, we derive the following estimate: Assuming that the mutants in the library only have variation in growth rate, the mutant library abundance should satisfy (11) xlib<ϵ⋅gwt|glib−gwt|, where ω is the desired threshold on the relative error (Fig. S23), gwt is the wild-type growth rate, and glib is the growth rate of the library as a whole (Sec. S16). In the case of our single-gene knockout library, Eq. (11) predicts a maximum library abundance of 24.6% based on the wild-type and library growth rate (gwt = 0.406, glib = 0.389) for a relative error ω = 1%. Figure 5D shows that this maximum library abundance keeps the relative error below 1% for high-fitness mutants (bright yellow), because they are dominated by growth rate effects, but fails for mutants close to neutrality because they have a tradeoff between growth rate and lag time, which the estimate in Eq. (11) neglects. It is possible to derive a more precise bound that keeps the relative error below 1% for all mutant genotypes (in this case xlib = 0.01%, see vertical line in Fig. 5D), but this requires prior knowledge of the trait covariation in the mutant library (Sec. S16).

Previous benchmark experiments show deviation of bulk fitness from pairwise competition

Earlier version: Our models predict that higher-order effects amplify pairwise fitness measured in bulk compared to the fitness in a pairwise competition experiment. To test this prediction, we reanalyze two published benchmark experiments of bulk fitness measurements. The first data set by van Opijnen et al. [103] measured pairwise (1:1) and bulk competitions (xlib ≈ 99.9%) of 46 single-gene knockouts in Streptococcus pneumoniae using logit-encoded fitness per-generation and neutral strains as the reference group. Figure 6A shows that the fitness in bulk deviates from the diagonal, and the absolute error has a positive correlation with the underlying true fitness from pairwise competition (Figure S24A; Pearson r = 0.40,p = 0.006), similar to the pattern observed in our simulations (compare Figure S24A with Figure 5C case III, and Figure S24B with Figure S23C case III).

Revised version: Our models predict that higher-order effects amplify pairwise fitness measured in bulk compared to the fitness in a pairwise competition experiment. To test this prediction, we reanalyze two published benchmark experiments of bulk fitness measurements. The first data set by van Opijnen et al. [106] measured pairwise (1:1) and bulk competitions (xlib ≈ 99.9%) of 46 single-gene knockouts in Streptococcus pneumoniae using logit-encoded fitness per-generation and neutral strains as the reference group. Figure 6A shows that the fitness in bulk deviates from the diagonal, and the absolute error has a positive correlation with the underlying true fitness from pairwise competition (Fig. S24A; Pearson r = 0.40,p = 0.006), similar to the pattern observed in our simulations (compare Fig. S24A with Fig. 5C case III, and Fig. S24B with Fig. S23C case III).

Empirical benchmarks of bulk fitness against pairwise competition experiments.

Earlier version: (A) Fitness measurements as reported by Opijnen et al. [103] for knockouts of S. pneumoniae. Pairwise fitness was measured in 1:1 competitions of knockouts strains with the wild-type over a single growth cycle using the per-generation statistic W=1+sgenlogit [26]. Bulk fitness was measured from a single growth cycle using 100% initial abundance of the mutant library and a set of neutral strains as the reference to calculate a per-generation fitness, replacing LFCwt with the library LFC in Eq. (S31) (Sec. S5). Since we were unable to obtain the original data, we extracted these fitness data from Figure 3c in the original publication and subtracted 1 to plot as sgenlogit. See Figure S24 for the absolute and relative error of bulk fitness based on these data. (B) Fitness measurements as reported by Levy et al. [13] for evolved genotypes of of S. cerevisiae after 80 generations of the evolution experiment. Pairwise fitness was measured in 1:1 competitions of knockouts trains with a fluorescently-labeled wild-type over seven growth cycle using relative fitness per-cycle under the log-encoding scyclelog, divided by a fixed number generations inferred from the dilution factor. Bulk fitness was measured from relative abundance data throughout the evolution experiment and using a multi-step algorithm, Bayesian inference, and a set of neutral strains as the reference, leading to a pairwise fitness estimate equivalent to scyclelogit (compare SI Eq. (50) in Levy et al. [13]). We extracted the fitness data from the table presented in SI Sec. 8.1 of the original publication. See Figure for the absolute and relative error of bulk fitness based on these data.

Revised version: (A) Fitness measurements as reported by Opijnen et al. [106] for knockouts of S. pneumoniae. Pairwise fitness was measured in 1:1 competitions of knockouts strains with the wild-type over a single growth cycle using the per-generation statistic W=1+sgenlogit [26]. Bulk fitness was measured from a single growth cycle using 100% initial abundance of the mutant library and a set of neutral strains as the reference to calculate a per-generation fitness, replacing LFCwt with the library LFC in Eq. (S31) (Sec. S5). Since we were unable to obtain the original data, we extracted these fitness data from Fig. 3c in the original publication and subtracted 1 to plot as sgenlogit. See Fig. S24 for the absolute and relative error of bulk fitness based on these data. (B) Fitness measurements as reported by Levy et al. [13] for evolved genotypes of of S. cerevisiae after 80 generations of the evolution experiment. Pairwise fitness was measured in 1:1 competitions of knockouts trains with a fluorescently-labeled wild-type over seven growth cycle using relative fitness per-cycle under the log-encoding scyclelog, divided by a fixed number generations inferred from the dilution factor. Bulk fitness was measured from relative abundance data throughout the evolution experiment and using a multi-step algorithm, Bayesian inference, and a set of neutral strains as the reference, leading to a pairwise fitness estimate equivalent to scyclelogit (compare SI Eq. (50) in Levy et al. [13]). We extracted the fitness data from the table presented in SI Sec. 8.1 of the original publication. See Fig. for the absolute and relative error of bulk fitness based on these data.

Earlier version: The second data set by Levy et al. [13] measured pairwise competitions (1:1) and relative abundances for 33 lab-evolved genotypes of S. cerevisiae with beneficial fitness. The bulk fitness is inferred under the log-encoding and per-cycle, with neutral strains as the reference group. Figure 6B shows that the fitness of strains in the bulk evolution environment is typically higher than in pair-wise competition, but the absolute error has no positive correlation (Figure S25A; Pearson r = −0.26,p = 0.140) and thus does not support the pattern we see our simulations (compare Figure S25A with Figure Figure case III, and S25B with Figure S23C case III).

Revised version: The second data set by Levy et al. [13] measured pair-wise competitions (1:1) and relative abundances for 33 lab-evolved genotypes of S. cerevisiae with beneficial fitness. The bulk fitness is inferred under the log-encoding and per-cycle, with neutral strains as the reference group. Figure 6B shows that the fitness of strains in the bulk evolution environment is typically higher than in pair-wise competition, but the absolute error has no positive correlation (Fig. S25A; Pearson r = −0.26,p = 0.140) and thus does not support the pattern we see our simulations (compare Fig. S25A with Fig. 5 case III, and Fig. S25B with Fig. S23C case III).

Discussion

Significance of fitness quantification choices for microbial ecology and evolution

In this work, we have introduced a conceptual framework to derive common statistics of relative fitness from three essential choices:

The choice of the state variable for relative abundance (encoding m).

The choice of the time scale for the change in relative abundance (Δt).

Earlier version: In bulk competition experiments, the choice of subpopulation that acts as the reference for relative abundance of a chosen mutant genotype (Figure 5B).

Revised version: In bulk competition experiments, the choice of sub-population that acts as the reference for relative abundance of a chosen mutant genotype (Fig. 5B).

Earlier version: The combination of these choices leads to a range of fitness statistics, including those commonly used in population genetics [80], experimental evolution [11, 12, 26], and gene knockout screens [17–19, 42, 43, 77, 78, 103].

Revised version: The combination of these choices leads to a range of fitness statistics, including those commonly used in population genetics [83], experimental evolution [11, 12, 26], and gene knockout screens [17–19, 42, 43, 80, 81, 106].

Earlier version: However, as we compare these statistics in simulated competition experiments, we find that the choice of time scale can lead to different mutant rankings (Figures 2B and 4A,B), confound apparent epistasis between mutations (Figure 3B), and alter the apparent long-term trend in fitness (Figure 3F), while the choice of the reference subpopulation can lead to a systematic shift in inferred fitness values (Figure 5C). The quantitative differences between relative fitness statistics affect our ability to make evolutionary predictions. In particular, the fixation probability of a spontaneous mutation depends on the magnitude of relative fitness compared to other time scales like mutation and drift [66, 67], and in microbial populations with clonal interference, it depends on the entire distribution of fitness effects [58, 104]. The choice of statistic may also affect the outcome of a significance test in the context of gene essentiality tests (e.g., using log or logit with the test statistic in Wetmore et al. [18]). Since these measurements often serve as a first screen to narrow down the investigation to the top set of genes [17, 40], a difference in mutant ranking means the investigation might miss relevant genes because of the choice of fitness statistic.

Revised version: However, as we compare these statistics in simulated competition experiments, we find that the choice of time scale can lead to different mutant rankings (Figs. 2B and 4A,B), confound apparent epistasis between mutations (Fig. 3B), and alter the apparent long-term trend in fitness (Fig. 3F), while the choice of the reference subpopulation can lead to a systematic shift in inferred fitness values (Fig. 5C). The quantitative differences between relative fitness statistics affect our ability to make evolutionary predictions. Fixation probabilities and other population genetics quantities depend on ratios of parameters defining the rates of selection, mutation, and genetic drift processes [66, 67]. These rates must be measured in the same units. For example, in batch cultures with serial transfer, genetic drift is set by the dilution bottle-neck and thus operates per cycle (e.g., 1/N per cycle, where N is the bottleneck population size). Thus quantifying a mutation’s relative fitness s (equivalent to the rate of selection) per generation and comparing to genetic drift per cycle (Ns) would incorrectly predict the fixation probability, especially in cases where per-generation and per-cycle fitness rank mutants differently. The choice of statistic may also affect the outcome of a significance test in the context of gene essentiality tests (e.g., using log or logit with the test statistic in Wetmore et al. [18]). Since these measurements often serve as a first screen to narrow down the investigation to the top set of genes [17, 40], a difference in mutant ranking means the investigation might miss relevant genes because of the choice of fitness statistic.

Best practices for quantifying mutant fitness in high-throughput experiments

Based on these insights, we recommend the following choices for quantifying relative fitness of a mutant:

Use the logit encoding of relative abundance, because under the null model of logistic dynamics this linearizes the trajectory of relative abundance. Only use a different encoding if alternate, non-logistic dynamics are known.

For batch growth or any experiment with non-steady state growth, use a fixed extrinsic time scale (e.g., a single growth cycle), rather than an intrinsic time scale (e.g., the number of generations), which introduces additional variation between competitions. Only quantify fitness per generation for experiments that are known to have steady-state growth (e.g., in a chemostat).

In bulk competition experiments, always use the pairwise fitness relative to the wild-type (by including a barcoded wild-type or by grouping neutral mutants into a virtual wild-type) because this more closely matches the relative fitness in pairwise competition.

To control the error from higher-order interactions, we recommend minimizing relative abundance of the mutant library as practically feasible. For a given error tolerance, Eq. (11) gives an estimate for the maximum library abundance.

Earlier version: Our recommendations agree with previous criticisms of relative fitness per-generation [5, 8, 9, 91] and total relative fitness [59, 105] but differ from the standard practice in high-throughput genetic screens [18, 19, 77] and many evolution experiments [33, 88], including the LTEE [11, 26]. The choice of the reference group is a well-recognized issue in high-throughput evolutionary studies, in which one first estimates a total relative fitness and then subtracts a correction based on the fitness of the wild-type [12, 50, 52] or a mean population fitness [13, 41, 42, 51, 105]. In contrast, we recommend to choose the reference group at the level of relative abundance (see Methods for a step-by-step demonstration). While researchers must ultimately choose how to quantify fitness based on the specific goals and conditions of their experiment, we believe that a more consistent adherence to the recommendations above would be valuable to the field as a whole so that fitness data from one study can be quantitatively compared to fitness from another.

Revised version: Our recommendations agree with previous criticisms of relative fitness per-generation [5, 8, 9, 94] and total relative fitness [59, 107] but differ from the standard practice in high-throughput genetic screens [18, 19, 80] and many evolution experiments [33, 91], including the LTEE [11, 26]. The choice of the reference group is a well-recognized issue in high-throughput evolutionary studies, in which one first estimates a total relative fitness and then subtracts a correction based on the fitness of the wild-type [12, 50, 52] or a mean population fitness [13, 41, 42, 51, 107]. In contrast, we recommend to choose the reference group at the level of relative abundance (see Methods for a step-by-step demonstration). While researchers must ultimately choose how to quantify fitness based on the specific goals and conditions of their experiment, we believe that a more consistent adherence to the recommendations above would be valuable to the field as a whole so that fitness data from one study can be quantitatively compared to fitness from another.

How general is the disagreement between fitness statistics in pairwise competition?

Earlier version: The discrepancy in mutant ranking for relative fitness per-cycle scyclelogit and per-generation sgenlogit requires positive covariation between the wild-type and mutant log fold-changes (LFCs) across a set of pairwise competition experiments. This condition is not dependent on any particular mechanism or model, although our minimal model of population dynamics (Figure 2A) demonstrates one way this covariation can occur: when mutants have diminished yields and are present at sufficiently-high abundances to affect wild-type LFCs (Figure S7).

Revised version: The discrepancy in mutant ranking for relative fitness per-cycle scyclelogit and per-generation sgenlogit requires positive covariation between the wild-type and mutant log fold-changes (LFCs) across a set of pairwise competition experiments. This condition is not dependent on any particular mechanism or model, although our minimal model of population dynamics (Fig. 2A) demonstrates one way this covariation can occur: when mutants have diminished yields and are present at sufficiently-high abundances to affect wild-type LFCs (Fig. S7).

Earlier version: Our model focused on the effect of resource competition on fitness discrepancies since that interaction is almost always present between microbial genotypes, providing a baseline expectation. Moreover, resource competition is likely to be the only interaction in high-throughput fitness measurements of mutants with only small genetic differences (e.g., point mutations or gene knockouts) relative to a wild-type. We expect that in cases with additional interactions between mutants and the wild-type, there will be further variation in wild-type LFCs across competitions and thus even more discrepancies betwen fitness statistics. For example, mutants that stimulate wild-type growth via the secretion of nutrients should create positive LFC covariation across mutants and thus be ranked differently per-cycle and per-generation. Indeed, our reanalysis of LTEE competition data shows that fitness per-cycle scyclelogit and per-generation sgenlogit (equivalent to the statistic W=1+sgenlogit used in the LTEE) disagree on the fitness ranking at almost every time point (Figure 4A) and in the first 50,000 generations overall (Figure 4B).

Revised version: Our model focused on the effect of resource competition on fitness discrepancies since that interaction is almost always present between microbial genotypes, providing a baseline expectation. Moreover, resource competition is likely to be the only interaction in high-throughput fitness measurements of mutants with only small genetic differences (e.g., point mutations or gene knockouts) relative to a wild-type. We expect that in cases with additional interactions between mutants and the wild-type, there will be further variation in wild-type LFCs across competitions and thus even more discrepancies betwen fitness statistics. For example, mutants that stimulate wild-type growth via the secretion of nutrients should create positive LFC covariation across mutants and thus be ranked differently per-cycle and per-generation. Indeed, our reanalysis of LTEE competition data shows that fitness per-cycle scyclelogit and pergeneration sgenlogit (equivalent to the statistic W=1+sgenlogit used in the LTEE) disagree on the fitness ranking at almost every time point (Fig. 4A) and in the first 50,000 generations overall (Fig. 4B).

How general is the deviation of bulk fitness statistics from pairwise competition experiments?

Earlier version: The offset in total relative fitness (between grey and orange points in Figure 5C), equivalent to the mean library fitness (Eq. (S88) in Sec. S13), is independent from the population dynamics (Figure 5B, Sec. S13). In contrast, the error between pairwise fitness in bulk competition and in pairwise competition (orange points in Figure 5C) is specific to the population dynamics studied here, although it points to more general mechanisms. Specifically, we expect pairwise fitness in bulk competition to deviate from pairwise competition as soon as the mutant library is able to influence the chemical environment of the competition, for example, when a large fraction of mutants fails to complete respiration of the main carbon source and instead turns to fermentation [106–108], with the release of pathway intermediates like acetate that indirectly lower the pH in the growth medium [109, 110].

Revised version: The offset in total relative fitness (between grey and orange points in Fig. 5C), equivalent to the mean library fitness (Eq. (S88) in Sec. S13), is independent from the population dynamics (Fig. 5B, Sec. S13). In contrast, the error between pairwise fitness in bulk competition and in pairwise competition (orange points in Fig. 5C) is specific to the population dynamics studied here, although it points to more general mechanisms. Specifically, we expect pairwise fitness in bulk competition to deviate from pairwise competition as soon as the mutant library is able to influence the chemical environment of the competition, for example, when a large fraction of mutants fails to complete respiration of the main carbon source and instead turns to fermentation [108–110], with the release of pathway intermediates like acetate that indirectly lower the pH in the growth medium [111, 112].

The choice of fitness statistic in the light of other sources of discrepancy

Earlier version: Besides the conceptual choices of fitness quantification discussed in this article, experimental limitations can create discrepancies between replicate fitness measurements of the same mutant [36, 50, 111]. For example, relative abundance measurements entail sampling uncertainties (when we sample liquid for colony counting or DNA sequencing) [44, 59, 60], copy number variation [18, 44], as well as PCR jackpots and sequencing read errors [13, 101, 102]. Furthermore, there will inevitably be some variation in initial condition between replicates and fluctuations during the fitness assay [36, 102, 111].

Revised version: Besides the conceptual choices of fitness quantification discussed in this article, experimental limitations can create discrepancies between replicate fitness measurements of the same mutant [36, 50, 113]. For example, relative abundance measurements entail sampling uncertainties (when we sample liquid for colony counting or DNA sequencing) [44, 59, 60], copy number variation [18, 44], as well as PCR jackpots and sequencing read errors [13, 104, 105]. Furthermore, there will inevitably be some variation in initial condition between replicates and fluctuations during the fitness assay [36, 105, 113].

Methods

Inferring growth traits for the single-gene knockout collection in yeast

Earlier version: We use a previously published data set [29] for the single-gene knockout collection in Saccharomyces cerevisiae [96], where the authors track growth of each genotype for 47 hours in monoculture, using microwell plates with defined growth medium. We download this growth curve data set from the PROPHECY database [112], choosing specifically the data set measured in Synthetic Defined medium. We last accessed the PROPHECY website (http://prophecy.lundberg.gu.se/) on March 30, 2020, but as of November 15, 2025 the website is no longer accessible. We have included a reformatted version of the data with our code repository (https://github.com/justuswfink/24FitnessQuantification). From the raw time series of optical density, the authors subtracted a background correction based on blank wells, applied another correction for nonlinearity at high optical densities, and smoothed the growth curve to remove electrical noise [29].

Revised version: We use a previously published data set [29] for the single-gene knockout collection in Saccharomyces cerevisiae [99], where the authors track growth of each genotype for 47 hours in monoculture, using microwell plates with defined growth medium. We download this growth curve data set from the PROPHECY database [114], choosing specifically the data set measured in Synthetic Defined medium. We last accessed the PROPHECY website (http://prophecy.lundberg.gu.se/) on March 30, 2020, but as of November 15, 2025 the website is no longer accessible. We have included a reformatted version of the data with our code repository (https://github.com/justuswfink/24FitnessQuantification). From the raw time series of optical density, the authors subtracted a background correction based on blank wells, applied another correction for nonlinearity at high optical densities, and smoothed the growth curve to remove electrical noise [29].

We start with the curves in the published data set (9951 curves) and apply further trimming and smoothing steps (Sec. S6). From this data we calculate the time series of instantaneous per-capita growth rate N−1dN/dt (where N is the optical density) and identify time windows where the rate is approximately constant, which we interpret as distinct growth phases (Sec. S6). We only include curves that have a single phase of constant exponential growth followed by a stationary phase of approximately zero growth (9424 curves).

Earlier version: We quantify the biomass yield, maximum growth rate, and the lag time from each remaining curve as follows. First, we estimate the initial abundance Ninitial (average optical density over first three time points) and the final abundance Nfinal (average optical density over stationary phase) from the growth curve. Then we calculate the biomass yield as Y = (Nfinal − Ninitial)/R(0), where R(0) = 111 mM (20 g/L) is the initial concentration of glucose (assuming glucose is the single limiting resource). To estimate the maximum growth rate, we average the instantaneous growth rate over the exponential phase. Finally, we estimate the lag time from the intersection of the log initial abundance (log N0) with the slope of the maximum growth rate during the exponential growth phase [81]. We excluded curves with negative initial OD, curves with negative inferred lag times, and curves with a low quality of fit between the measured time series and a simulated curve based on the inferred trait values (R2 < 0.95; see below for the model). Our final database includes trait estimates for 9195 curves, which represents 92.4% of the original data set [29].

Revised version: We quantify the biomass yield, maximum growth rate, and the lag time from each remaining curve as follows. First, we estimate the initial abundance Ninitial (average optical density over first three time points) and the final abundance Nfinal (average optical density over stationary phase) from the growth curve. Then we calculate the biomass yield as Y = (Nfinal − Ninitial)/R(0), where R(0) = 111 mM (20 g/L) is the initial concentration of glucose (assuming glucose is the single limiting resource). To estimate the maximum growth rate, we average the instantaneous growth rate over the exponential phase. Finally, we estimate the lag time from the intersection of the log initial abundance (log N0) with the slope of the maximum growth rate during the exponential growth phase [84]. We excluded curves with negative initial OD, curves with negative inferred lag times, and curves with a low quality of fit between the measured time series and a simulated curve based on the inferred trait values (R2 < 0.95; see below for the model). Our final database includes trait estimates for 9195 curves, which represents 92.4% of the original data set [29].

Earlier version: Each single-gene deletion strain was measured in two technical replicate growth curves, using a second plate with identical layout in the same plate reader. Some genotypes have only one estimate of the growth traits because the other replicate did not pass our filters (273 genotypes), but for most genotypes we retain two replicate estimates in our final data set (4163 genotypes); a few genotypes even have three (2 genotypes) or four replicate estimates (54 genotypes) because these geno-types were included multiple times by the original authors [29]. From these replicates we finally calculate the average yield, growth rate, and lag time for each geno-type. The data set also contained many replicate growth curves of the wild-type strain, 374 of which passed our filters. Since wild-type traits inferred from these replicates had large variation (potentially due to measuring the wild-type across many different plates and days), we define the wild-type trait as the median, not the mean, of these replicates. Although the wild-type measurements included in this data set have large trait variation (orange dots in Figure S4B,C), the fact that the variation across knockouts is significantly greater than for the wild-type replicates in terms of growth rate (Levene’s test p = 1.3 × 10−9) and lag time (Levene’s test p = 0.025), and that the replicate measurements of gene knockouts are correlated (Figure S5), suggests that these trait measurements do capture true genetic variation. In contrast, the variation in biomass yield across knockouts is not significantly different from the variation across wild-type replicates (Levene’s test p = 0.44), suggesting that the variation in biomass yield is driven by non-genetic factors like the abundance or batch of resources used for each growth measurement.

Revised version: Each single-gene deletion strain was measured in two technical replicate growth curves, using a second plate with identical layout in the same plate reader. Some genotypes have only one estimate of the growth traits because the other replicate did not pass our filters (273 genotypes), but for most genotypes we retain two replicate estimates in our final data set (4163 genotypes); a few genotypes even have three (2 genotypes) or four replicate estimates (54 genotypes) because these geno-types were included multiple times by the original authors [29]. From these replicates we finally calculate the average yield, growth rate, and lag time for each geno-type. The data set also contained many replicate growth curves of the wild-type strain, 374 of which passed our filters. Since wild-type traits inferred from these replicates had large variation (potentially due to measuring the wild-type across many different plates and days), we define the wild-type trait as the median, not the mean, of these replicates. Although the wild-type measurements included in this data set have large trait variation (orange dots in Fig. S4B,C), the fact that the variation across knockouts is significantly greater than for the wild-type replicates in terms of growth rate (Levene’s test p = 1.3 × 10−9) and lag time (Levene’s test p = 0.025), and that the replicate measurements of gene knockouts are correlated (Fig. S5), suggests that these trait measurements do capture true genetic variation. In contrast, the variation in biomass yield across knockouts is not significantly different from the variation across wild-type replicates (Levene’s test p = 0.44), suggesting that the variation in biomass yield is driven by non-genetic factors like the abundance or batch of resources used for each growth measurement.

Simulating population dynamics in competition experiments

Earlier version: Throughout this work, we model the competition between a set of genotypes during a single batch culture growth cycle using the following model of population dynamics [94, 95]: (12a) 1NjdNjdt=gj(t,R),Nj(0)=xj(0)N0 (12b) dRdt=−∑j1YjdNjdt,R(0)=R0 where for each genotype j, Nj is the absolute abundance, xj(0) the initial relative abundance, gj(t, R) the growth rate, and Yj the biomass yield, where the genotypes include a wild-type as well as one or more mutants. The initial absolute abundance of all genotypes together is N0, and the concentration of the single limiting resource is R with initial value R0. Note that Eq. (12b) assumes that cells consume resources only for biomass growth and not for maintenance of existing biomass.

Revised version: Throughout this work, we model the competition between a set of genotypes during a single batch culture growth cycle using the following model of population dynamics [97, 98]: (12a) 1NjdNjdt=gj(t,R),Nj(0)=xj(0)N0 (12b) dRdt=−∑j1YjdNjdt,R(0)=R0 where for each genotype j, Nj is the absolute abundance, xj(0) the initial relative abundance, gj(t, R) the growth rate, and Yj the biomass yield, where the genotypes include a wild-type as well as one or more mutants. The initial absolute abundance of all genotypes together is N0, and the concentration of the single limiting resource is R with initial value R0. Note that Eq. (12b) assumes that cells consume resources only for biomass growth and not for maintenance of existing biomass.

Earlier version: To capture the sigmoidal shape of typical growth curves with a lag phase, exponential growth phase, and saturation phase (e.g., Figure S4A), we model the dependence of growth rate on time and resource concentration as (13) gj(t,R)=gjΘ(t−λj)Θ(R), where Θ is the step function: (14) Θ(z)={0z≤0,1z>0. In this model, a genotype has an initial lag phase of time λj where no growth occurs, followed by a phase of constant exponential growth at rate gj, and ending when the resource concentration R reaches zero. The time to resource depletion, which we also call the saturation time, is defined by the implicit equation R(tsat) = 0. The simple form of the growth response (Eq. (13)) means that the absolute abundance of genotype j at saturation is (15) Nj(tsat)=Nj(t0)egj(tsat−λj)Θ(tsat−λj) and its log fold-change is (16) LFCj=gj(tsat−λj)Θ(tsat−λj). To calculate relative fitness per-cycle (Eq. (9) and per-generation (Eq. (10)), we numerically determine the saturation time tsat by integrating the biomass and resource dynamics (Eq. (12)) as described in previous work [94, 95]. We note that it also possible to derive an approximate expression for the saturation time tsat which we use for our analytical calculations (Sec. S7).

Revised version: To capture the sigmoidal shape of typical growth curves with a lag phase, exponential growth phase, and saturation phase (e.g., Fig. S4A), we model the dependence of growth rate on time and resource concentration as (13) gj(t,R)=gjΘ(t−λj)Θ(R), where Θ is the step function: (14) Θ(z)={0z≤01z>0. In this model, a genotype has an initial lag phase of time εj where no growth occurs, followed by a phase of constant exponential growth at rate gj, and ending when the resource concentration R reaches zero. The time to re-source depletion, which we also call the saturation time, is defined by the implicit equation R(tsat) = 0. The simple form of the growth response (Eq. (13)) means that the absolute abundance of genotype j at saturation is (15) Nj(tsat)=Nj(t0)egj(tsat−λj)Θ(tsat−λj) and its log fold-change is (16) LFCj=gj(tsat−λj)Θ(tsat−λj). To calculate relative fitness per-cycle (Eq. (9) and per-generation (Eq. (10)), we numerically determine the saturation time tsat by integrating the biomass and resource dynamics (Eq. (12)) as described in previous work [97, 98]. We note that it also possible to derive an approximate expression for the saturation time tsat which we use for our analytical calculations (Sec. S7).

Testing epistasis between mutations in population dynamics model

Earlier version: Consider a wild-type strain with growth rate gwt = 1, lag time λwt = 2, and biomass yield Ywt = 1. We define six basic mutation that represent the six directions of the trait space in our model: a mutation that increases growth rate (gmut = 1.2; symbol +g), a mutation that decreases growth rate (gmut = 0.8; symbol −g), a mutation that increases lag time (λmut = 3; symbol +λ), a mutation that decreases lag time (λmut = 1; symbol −λ), a mutation that increases yield (Ymut = 4; symbol +Y ), and a mutation that decreases yield (Ymut = 0.25; symbol −Y ). We combine these mutations into double mutants as follows: mutations are additive at the level of traits (e.g., the double mutant (+g, −g) is neutral) and mutations in different traits are independent. For a given pair of mutations, epistasis is defined as [97] (17) E=s(A,B)−s(A)−s(B), where s(A) is the relative fitness of the single mutant with mutation A, s(B) is the relative fitness of the single mutant with mutation B, and s(A, B) is the relative fitness of the double mutant with both A and B. We estimate the fitness of each mutant by simulating one growth cycle of competition with the wild-type (at initially equal abundance) using an initial biomass of N0 = 0.01 and initial resource concentration R0 = 0.5. These conditions are chosen such that mutations that increase biomass yield reach a maximum density of up to N (tsat) = 2, on the same scale as the optical density measured for the empirical knockouts of S. cerevisiaea (Figure S4C).

Revised version: Consider a wild-type strain with growth rate gwt = 1, lag time εwt = 2, and biomass yield Ywt = 1. We define six basic mutation that represent the six directions of the trait space in our model: a mutation that increases growth rate (gmut = 1.2; symbol +g), a mutation that decreases growth rate (gmut = 0.8; symbol − g), a mutation that increases lag time (εmut = 3; symbol +ε), a mutation that decreases lag time (εmut = 1; symbol − ε), a mutation that increases yield (Ymut = 4; symbol +Y), and a mutation that decreases yield (Ymut = 0.25; symbol −Y). We combine these mutations into double mutants as follows: mutations are additive at the level of traits (e.g., the double mutant (+g, −g) is neutral) and mutations in different traits are independent. For a given pair of mutations, epistasis is defined as [100] (17) E=s(A,B)−s(A)−s(B), where s(A) is the relative fitness of the single mutant with mutation A, s(B) is the relative fitness of the single mutant with mutation B, and s(A, B) is the relative fitness of the double mutant with both A and B. We estimate the fitness of each mutant by simulating one growth cycle of competition with the wild-type (at initially equal abundance) using an initial biomass of N0 = 0.01 and initial resource concentration R0 = 0.5. These conditions are chosen such that mutations that increase biomass yield reach a maximum density of up to N (tsat) = 2, on the same scale as the optical density measured for the empirical knockouts of S. cerevisiaea (Fig. S4C).

Quantifying total and pairwise relative fitness in bulk competition experiments

Earlier version: We define total and pairwise relative fitness,two alternative measures of fitness in bulk competition experiments, as follows. Consider a population with multiple genotypes such that each genotype i has absolute abundance Ni(t) at time t. The total relative abundance of the genotype is (18) xi(t)=Ni(t)∑jNj(t). The “total” here refers to the fact that this is the abundance of genotype i relative to all other genotypes in the population. Similarly, we can think of the definition of relative fitness discussed in the main text (Eq. (4)) as the total relative fitness of this genotype under an encoding (19) σim=ddtm(xi(t)). However, sometimes we want to track the dynamics of a genotype relative to another specific genotype; for example, to follow the fate of a mutant against the wild-type in a bulk competition experiment. For a pair of genotypes i ≠ j, we thus define the pairwise relative abundance (20) xij(t)=xi(t)xi(t)+xj(t), which is the relative abundance of genotype i in a subpopulation of only genotypes i and j. To predict the change in the pairwise relative abundance, we can use the slope (21) sijm=ddtm(xij(t)), which defines the pairwise relative fitness of genotype i with respect to genotype j (see Sec. for a generalize definition that covers the case i = j). In the special case of a population with only two strains, the pair-wise relative fitness and total relative fitness are identical (σ2m=s21m) but they may differ with more than two geno-types (Sec. S13).

Revised version: We define total and pairwise relative fitness,two alternative measures of fitness in bulk competition experiments, as follows. Consider a population with multiple genotypes such that each genotype i has absolute abundance Ni(t) at time t. The total relative abundance of the genotype is (18) xi(t)=Ni(t)∑jNj(t). The “total” here refers to the fact that this is the abundance of genotype i relative to all other genotypes in the population. Similarly, we can think of the definition of relative fitness discussed in the main text (Eq. (4)) as the total relative fitness of this genotype under an encoding m: (19) σim=ddtm(xi(t)). However, sometimes we want to track the dynamics of a genotype relative to another specific genotype; for example, to follow the fate of a mutant against the wild-type in a bulk competition experiment. For a pair of genotypes i ≠ j, we thus define the pairwise relative abundance (20) xij(t)=xi(t)xi(t)+xj(t), which is the relative abundance of genotype i in a sub-population of only genotypes i and j. To predict the change in the pairwise relative abundance, we can use the slope (21) sijm=ddtm(xij(t)), which defines the pairwise relative fitness of genotype I with respect to genotype j (see Sec. for a generalize definition that covers the case i = j). In the special case of a population with only two strains, the pair-wise relative fitness and total relative fitness are identical (σ2m=s21m) but they may differ with more than two genotypes (Sec. S13).

Earlier version: It is also possible to define the total and pairwise relative fitness statistics as finite differences over a time interval, rather than as instantaneous derivatives. For example, we can define them over a growth cycle starting at t = 0 and ending at t = tsat: (22) σi,cyclem=m(xi(tsat))−m(xi(0)), (23) sij,cyclem=m(xij(tsat))−m(xij(0)). For these fitness statistics in the bulk competitions, we use the logit encoding (24) σi,cyclelogit=log⁡xi(tsat)1−xi(tsat)−log⁡xi(0)1−xi(0), (25) sij,cyclelogit=log⁡xi(tsat)xi(0)−log⁡xj(tsat)xj(0), where we have rewritten sij,cyclelogit in the form that it is presented in bulk competition experiments [15, 51, 52]. The logit encoding has mathematical advantages for coarse-graining the relative fitness of genotype groups (Sec. S11) but using the log encoding is another common choice in the literature [17, 41, 50]. However, the relative abundances of the individual mutants in our simulated bulk competition experiments are low enough that these two encodings are approximately equivalent (logit x ≈ log x for x ≈ 1).

Revised version: It is also possible to define the total and pairwise relative fitness statistics as finite differences over a time interval, rather than as instantaneous derivatives. For example, we can define them over a growth cycle starting at t = 0 and ending at t = tsat: (22) σi,cyclem=m(xi(tsat))−m(xi(0)), (23) sij,cyclem=m(xij(tsat))−m(xij(0)). For these fitness statistics in the bulk competitions, we use the logit encoding (24) σi,cyclelogit=log⁡xi(tsat)1−xi(tsat)−log⁡xi(0)1−xi(0), (25) sij,cyclelogit=log⁡xi(tsat)xi(0)−log⁡xj(tsat)xj(0), where we have rewritten sij,cycle in the form that it is presented in bulk competition experiments [15, 51, 52]. The logit encoding has mathematical advantages for coarse-graining the relative fitness of genotype groups (Sec. S11) but using the log encoding is another common choice in the literature [17, 41, 50]. However, the relative abundances of the individual mutants in our simulated bulk competition experiments are low enough that these two encodings are approximately equivalent (logit x ≈ log x for x ≪ 1).

Earlier version: Unlike the comparison between per-cycle and per-generation relative fitness where we focused on rank differences (Figure 2B), here we can evaluate absolute differences in fitness estimates because the total and pairwise fitness in bulk fitness are measured in the same units as the relative fitness in pairwise competition. For all mutant genotypes, we calculate the absolute error between these bulk fitness estimates (Eqs. (24) and (25)) and the relative fitness in pairwise competitions (Eq. (9)), which we take as the ground truth (Figure 5C). Note that the pairwise competition could depend on the initial relative abundance of the mutant; we have chosen a very low relative abundance (10−6) that mimics a mutant arising de novo and where this dependence is very weak.

Revised version: Unlike the comparison between per-cycle and pergeneration relative fitness where we focused on rank differences (Fig. 2B), here we can evaluate absolute differences in fitness estimates because the total and pairwise fitness in bulk fitness are measured in the same units as the relative fitness in pairwise competition. For all mu-tant genotypes, we calculate the absolute error between these bulk fitness estimates (Eqs. (24) and (25)) and the relative fitness in pairwise competitions (Eq. (9)), which we take as the ground truth (Fig. 5C). Note that the pairwise competition could depend on the initial relative abundance of the mutant; we have chosen a very low relative abundance (10−6) that mimics a mutant arising de novo and where this dependence is very weak.

Earlier version: Clearly, minimizing the abundance of the mutant library minimizes the strength of higher-order interactions (also reducing ranking disagreement, Figure S21B,C). As low mutant abundances create practical problems for experiments, it is valuable to identify a maximum abundance for the mutant library that keeps the error from bulk fitness estimates below a desired threshold (Eq. (11)). It is convenient to express this threshold in relative rather than absolute terms (compare Figure S23 to Figure 5C).

Revised version: Clearly, minimizing the abundance of the mutant library minimizes the strength of higher-order interactions (also reducing ranking disagreement, Fig. S21B,C). As low mutant abundances create practical problems for experiments, it is valuable to identify a maximum abundance for the mutant library that keeps the error from bulk fitness estimates below a desired threshold (Eq. (11)). It is convenient to express this threshold in relative rather than absolute terms (compare Fig. S23 to Fig. 5C).

A guide to calculate pairwise relative fitness under the logit encoding from bulk competition data

Here ri(t) represents the best estimate of read numbers for lineage i after correcting for experimental artifacts, like jackpot events in the sequencing (e.g., see discussion in Levy et al. [13]).

Calculate the total relative abundance xi of each lineage i by dividing with the total number of reads: (26) xi(t)=ri(t)∑i=1nri(t), where the sum is over all n lineages in the experiment.

Define a group of lineages i = 1, …, k as the reference (e.g., a group of neutral mutants) and sum their total relative abundance (27) xref(t)=∑i=1kxi(t).

Calculate the pairwise relative abundance xi,ref of each lineage by (28) xi,ref(t)=xi(t)xi(t)+xref(t).

Calculate the pairwise fitness per-cycle for each strain i under the logit encoding (29) si,reflogit=1Δt[log⁡xi,ref(t2)1−xi,ref(t2)−log⁡xi,ref(t1)1−xi,ref(t1)], where Δt is the number of growth cycles between time points t1 and t2.